---
title: 'Working with DATE/TIME Functions and Operators'
description: 'Chapter description goes here.'
---

## Functions for manipulating date/time data

```yaml
type: VideoExercise
key: 11fcc5db9e
xp: 50
```

`@projector_key`
57f28aba5697c148d36c7cb7dd45b7b3

---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
